```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import sympy as sym
```

```python
import PHYS121
PHYS121.Installer()
```

Output:
```
<IPython.core.display.HTML object>
```

```python
# Import Otter for the purposes of autograding
import otter

# Initialize Otter (for the purposes of auto-grading some cells)
grader = otter.Notebook("dataLogger.ipynb")
```

```python
z = sym.symbols('z')
```

***
**<span style="color:blue">Question 1</span>** **<span style="color:red">(1 mark)</span>**  
Enter a string.  The correct answer is 'hello world'.

**<span style="color:blue">Answer 1</span>**  
Replace the ... in the cell below with your answer.    
*** Please do not change anything to the left of the equals sign. ***

```python
text = 'hello world' # SOLUTION
```

```python
grader.check("q1")
```

***
**<span style="color:blue">Question 2</span>** **<span style="color:red">(1 mark)</span>**  
Enter a float.  The correct answer is 2.34.

**<span style="color:blue">Answer 2</span>**  
Replace the ... in the cell below with your answer.    
*** Please do not change anything to the left of the equals sign. ***

```python
num = 2.34 # SOLUTION
```

```python
grader.check("q2")
```

***
Just for fun, let's import and display the PHYS121_DataLogger.txt file using a Pandas dataframe.

```python
df = pd.read_csv('PHYS121_DataLogger.txt', sep='\t')
df
```

***
Once you've completed this notebook:
- Save your work.
- Run 'grader.check_all()' to confirm that you've completed all required tasks.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file.
- Upload the .zip file to the PHYS 121 Canvas gradebook.
- **Do NOT change the number of the .zip file.**
- **Do NOT modify the contents of the .zip file.**

Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(files = ['PHYS121_DataLogger.txt'], pdf = False)
```

***
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/Digital-Assessment-Techniques-CoP/main/images/ubc-logo-full.jpg" width="500"/>

Last update: January 21, 2024

Tests for q1 (points: None)
```
>>> isinstance(text, str)
True
```
failure_message: text must be a string (enclosed in single or double quotes).
```
>>> test = True
>>> if isinstance(text, str):
...     test = text == "hello world"
>>> test
True
```
failure_message: Double check the string you entered.


Tests for q2 (points: None)
```
>>> isinstance(num, float)
True
```
failure_message: num must be a float (decimal number).
```
>>> test = True
>>> if isinstance(num, float):
...     test = num == 2.34
>>> test
True
```
failure_message: Double check the float you entered.
